```python

```

# CSCI 3155: Assignment 9

Topics: 
- Type Checking
- Basics of Objects
- Generics

Readings: Notes posted on moodle.  Chapters from Odersky

- Chapter 4 (revise)
- Chapter 10
- Chapter 11 (take note of Any, AnyVal, Null, Nothing)
- Chapter 12 
- Chapter 19

__Name__: Ben Jacobs

```python
// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
passed: (points: Int)Unit
```

## Problem 1 (25 points)

In this problem, we will extend our type system for handling references.  Here is a stripped down version of Lettuce that involves references.

$$\begin{array}{rcl}
\mathbf{Expr} & \rightarrow & Const(\mathbf{Double}) \\
& | & Ident(\mathbf{Identifier})\\
& | & Plus(\mathbf{Expr}, \mathbf{Expr})\\
& | & Div(\mathbf{Expr}, \mathbf{Expr}) \\
& | & Geq(\mathbf{Expr}, \mathbf{Expr}) \\
& | & And(\mathbf{Expr}, \mathbf{Expr}) \\
& | & IfThenElse(\mathbf{Expr}, \mathbf{Expr}, \mathbf{Expr})\\
& | & Let(\mathbf{Identifier}, \color{red}{\mathbf{Type}}, \mathbf{Expr}, \mathbf{Expr})\\
& | & FunDef(\mathbf{Identifier}, \color{red}{\mathbf{Type}}, \mathbf{Expr}) \\
& | & FunCall(\mathbf{Expr}, \mathbf{Expr}) \\
& | & \color{blue}{NewRef(\mathbf{Expr})} \\
& | & \color{blue}{DeRef(\mathbf{Expr})} \\
& | & \color{blue}{AssignRef(\mathbf{Expr}, \mathbf{Expr})} \\
\end{array}$$

We have a type system 
$$\begin{array}{rcl}
\mathbf{Type} & \rightarrow & NumType & \leftarrow \text{written as } \ num \\
& | & BoolType & \leftarrow \text{written as } \ bool \\
& | & FunType(\mathbf{Type} , \mathbf{Type} ) & \leftarrow \text{written as } \ t_1 \Rightarrow t_2 \\
& | & \color{red}{RefType(\mathbf{Type})} & \leftarrow \text{written as } \ ref(t_1) \\
\end{array}$$


### A (5 points)
Copy the program below to the next cell and add the missing type annotations labeled type1, type2, and type3 to the the program.


~~~
let x: type1 = 25 in 
    let y : type2 = NewRef( - x ) in 
     let f : type3 = NewRef( function (z: ref(num)) DeRef(z) - x >= 45 ) in 
          let z : ref(bool) = DeRef(f)(y)
~~~

<!-- grade_id: cell-128df73da0a393cd, points: 5, grade: true, solution: true -->
let x: num = 25 in 
    let y : ref(t1) = NewRef( - x ) in 
     let f : t1 => t2 = NewRef( function (z: ref(num)) DeRef(z) - x >= 45 ) in 
          let z : ref(bool) = DeRef(f)(y)

### B (10 points): Let us write some inference rules for working with ref types.
$$\newcommand\typeOf{\mathbf{typeOf}}$$
$$\newcommand\semRule[3]{\begin{array}{c} #1 \\ \hline #2 \\ \end{array}\;(\text{#3}) }$$

Recall from notes on "Types and Type Checking" that $\typeOf(\texttt{e}, \alpha)$ is the type of an 
expression $\texttt{e}$ under type environment $\alpha$. The type environment maps identifiers in the current scope to  their annotated types.


Let us write a rule for NewRef.

$$\semRule{ \typeOf(\texttt{e}, \alpha) = t,\; t \not= \mathbf{typeerror} }{\typeOf(\texttt{NewRef(e)}, \alpha) = ref(t) }{newref-ok}$$

It says that if $\texttt{e}$ receives type $t$ under type environment $\alpha$ and it is not a type error, then $\texttt{NewRef(e)}$ must receive the type $ref(t)$ under $\alpha$.


(i) Complete the missing terms for the rule for `DeRef` OK rule.
$$\semRule{\typeOf(\texttt{e}, \alpha) = ref(t)}{\typeOf(\texttt{DeRef(e)}, \alpha) = \color{red}{???_1}}{deref-ok}$$

(ii) Complete the missing terms for the rule for `DeRef` error.
$$\semRule{\typeOf(\texttt{e}, \alpha) = t,\ t \not= \color{red}{???_2} }{\typeOf(\texttt{DeRef(e)}, \alpha) = \mathbf{typeerror}}{deref-nok}$$

(iii) Complete the missing terms for `AssignRef` OK rule.
$$\semRule{\typeOf(\texttt{e1}, \alpha) = \color{red}{???_3}, \typeOf(\texttt{e2}, \alpha) = t}{\typeOf(\texttt{AssignRef(e1, e2)}, \alpha) =   \color{red}{???_4}}{assignref-ok}$$


Write your answers in the next cell.

<!-- grade_id: cell-dd5541e0eaef1e6f, points: 10, grade: true, solution: true -->
1: t
2: ref(t)
3: t1
4: t1 = t

### C (10 points): Write the type checker with references by filling in the missing code.

<!-- grade_id: cell-a9e5e31b52e7ee1b, points: 0, grade: false, solution: true -->
```python
sealed trait Type
case object NumType extends Type
case object BoolType extends Type
case class FunType(t1: Type, t2: Type) extends Type
// TODO: Write a new case class for RefType
case class RefType(t1: Type) extends Type

sealed trait Program
sealed trait Expr

case class Const(f: Double) extends Expr
case class Ident(s: String) extends Expr
case class Plus(e1: Expr, e2: Expr) extends Expr
case class Minus(e1: Expr, e2: Expr) extends Expr
case class Geq(e1: Expr, e2: Expr) extends Expr
case class IfThenElse(e1: Expr, e2: Expr, e3: Expr) extends Expr
case class Let(x: String, xType: Type, e1: Expr, e2: Expr) extends Expr
case class FunDef(id: String, idType: Type, e: Expr) extends Expr
case class FunCall(calledFun: Expr, argExpr: Expr) extends Expr
case class NewRef(e: Expr) extends Expr
case class DeRef(e: Expr) extends Expr
case class AssignRef(e1: Expr, e2: Expr) extends Expr
case class TopLevel(e: Expr) extends Program

def typeEquals(t1: Type, t2: Type): Boolean = t1 == t2
case class TypeErrorException(s: String) extends Exception

```

Output:
```
defined trait Type
defined object NumType
defined object BoolType
defined class FunType
defined class RefType
defined trait Program
defined trait Expr
defined class Const
defined class Ident
defined class Plus
defined class Minus
defined class Geq
defined class IfThenElse
defined class Let
defined class FunDef
defined class FunCall
defined class NewRef
defined class DeRef
defined class AssignRef
defined class TopLevel
typeEquals: (t1: Type, t2: Type)Boolean
defined class TypeErrorException
```

<!-- grade_id: cell-809dce336587273e, points: 0, grade: false, solution: true -->
```python
def typeOf(e: Expr, alpha: Map[String, Type]): Type = {
    def checkType(opName: String, e1: Expr, t1: Type, e2: Expr, t2: Type, resType: Type): Type = {
        val t1hat = typeOf(e1, alpha)
        if (! typeEquals(t1hat, t1)){
            throw new TypeErrorException(s"Type mismatch in arithmetic/comparison/bool op $opName, Expected type $t1, obtained $t1hat")
        }
        
        val t2hat = typeOf(e2, alpha)
        if (! typeEquals(t2hat, t2)){
            throw new TypeErrorException(s"Type mismatch in arithmetic/comparison/bool op $opName, Expected type $t2, obtained $t2hat")
        }
        
        resType
    }
    
    e match {
        case Const(f) => NumType
        case Ident(s) => {if (alpha contains s)
                             alpha(s)
                          else 
                             throw TypeErrorException(s"Unknown identifier $s")}
        case Plus(e1, e2) =>  checkType("Plus", e1,  NumType, e2, NumType, NumType)
        case Minus(e1, e2) => checkType("Minus",e1,  NumType, e2, NumType, NumType)
        case Geq(e1, e2) => checkType("Geq", e1,  NumType, e2, NumType, BoolType)
        case IfThenElse(e, e1, e2) => {
            val t = typeOf(e, alpha)
            if (t == BoolType){
                val t1 = typeOf(e1, alpha)
                val t2 = typeOf(e2, alpha)
                if (typeEquals(t1, t2))
                    t1
                else 
                    throw TypeErrorException(s"If then else returns unequal types $t1 and $t2")
            } else {
                throw TypeErrorException(s"If then else condition expression not boolean $t")
            }
        }

        case Let(x, t, e1, e2) => {
            val t1 = typeOf(e1, alpha)
            if (typeEquals(t1, t)){
                val newAlpha = alpha + (x -> t)
                typeOf(e2, newAlpha)
            } else {
                throw TypeErrorException(s"Let binding has type $t whereas it is bound to expression of type $t1")
            }
        }

        case FunDef(x, t1, e) => {
            val newAlpha = alpha + (x -> t1)
            val t2 = typeOf(e, newAlpha)
            FunType(t1, t2)
        }

        case FunCall(e1, e2) => {
            val ftype = typeOf(e1, alpha)
            ftype match {
                case FunType(t1, t2) => {
                    val argType = typeOf(e2, alpha)
                    if (typeEquals(argType, t1)){
                        t2
                    } else {
                        throw TypeErrorException(s"Call to function with incompatible argument type. Expected $t1, obtained $argType")
                    }
                }
                case _ => { throw TypeErrorException(s"Call to function but with a non function type $ftype")}

            }
        }

        case NewRef(e) => {
            val t = typeOf(e,alpha);
            t
        }
        
        case AssignRef(e1, e2) => {
           val t = typeOf(e,alpha);
            t
        }
        
        case DeRef(e) => {
            val t = typeOf(e,alpha);
            t
        }
        
    }
}

def typeOfProgram(p: Program) = p match {
    case TopLevel(e) => {
            val t = typeOf(e, Map())
            println(s"Program type computed successfully as $t")
            t
    }
}
```

Output:
```
typeOf: (e: Expr, alpha: Map[String,Type])Type
typeOfProgram: (p: Program)Type
```

<!-- grade_id: cell-e6bbf61c650e47bb, points: 2, grade: true, solution: false -->
```python
//BEGIN TEST

/* 
let x : ref(num) = NewRef(10 ) in 
   let dummy: num = AssignRef(x, 30) in 
       DeRef(x)
       */

val p1 = Let("x", RefType(NumType), NewRef(Const(10)), Let("dummy", NumType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t1 = typeOfProgram(TopLevel(p1))
assert(t1 == NumType, "Test 1 failed: answer should be NumType")
passed(2)
//END TEST
```

Output:
```
at typeOf(<console>:171)
  at typeOfProgram(<console>:217)
  ... 44 elided
```

<!-- grade_id: cell-e80a253dd4c598ac, points: 2, grade: true, solution: false -->
```python
//BEGIN TEST
/* 
let x : ref(num) = NewRef(function(z: num) z + 10) in 
   let dummy: num = AssignRef(x, 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p2 = Let("x", RefType(NumType), NewRef(fdef), Let("dummy", NumType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t2 = try {
   typeOfProgram(TopLevel(p2))
   assert(false, "The program should not receive a type")
} catch {
    case TypeErrorException(msg) => s"OK -- caught a type error exception: $msg"
    case e => print(e); assert(false, "Please throw TypeErrorException(message) when a type failure occurs")
}
passed(2)
//END TEST
```

Output:
```
*** Tests Passed (2 points) ***
null
```

<!-- grade_id: cell-f6584c1da70c0ae8, points: 2, grade: true, solution: false -->
```python
//BEGIN TEST
/* 
let x : ref(num => num) = NewRef(function(z: num) z + 10) in 
   let dummy: num = AssignRef(x, 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p3 = Let("x", RefType(FunType(NumType, NumType)), NewRef(fdef), Let("dummy", NumType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t3 = try {
   typeOfProgram(TopLevel(p3))
   assert(false, "The program should not receive a type")
} catch {
    case TypeErrorException(msg) => s"OK -- caught a type error exception: $msg"
    case e => print(e); assert(false, "Please throw TypeErrorException(message) when a type failure occurs")
}
passed(2)
//END TEST
```

Output:
```
*** Tests Passed (2 points) ***
null
```

<!-- grade_id: cell-9060540ef32281bd, points: 2, grade: true, solution: false -->
```python
//BEGIN TEST
/* 
let x : ref(num => num) = NewRef(function(z: num) z + 10) in 
   let dummy: num = AssignRef(NewRef(35), 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p4 = Let("x", RefType(FunType(NumType, NumType)), NewRef(fdef), Let("dummy", NumType, AssignRef(NewRef(Const(35)), Const(30) ), DeRef(Ident("x"))) )
val t4 =  typeOfProgram(TopLevel(p4))
assert(t4 == FunType(NumType, NumType), "Test failed")
passed(2)
//END TEST
```

Output:
```
at typeOf(<console>:171)
  at typeOfProgram(<console>:217)
  ... 44 elided
```

<!-- grade_id: cell-b63b128c12f0247d, points: 2, grade: true, solution: false -->
```python
//BEGIN TEST

val y = Ident("y")
val x = Ident("x")
val z = Ident("z")
val t1 = RefType(NumType)
val t2 = FunType(t1, NumType)
val t3 = RefType(t2)

val v1 = FunCall(DeRef(y), x)
val v2 = Let("z", RefType(BoolType), NewRef(Geq(DeRef(x), Const(25))), v1)
val fdef = FunDef("z", RefType(NumType), Plus(DeRef(z), DeRef(x)))
val v3 = Let("y", t3, NewRef(fdef), v2)
val v4 = Let("x", t1, NewRef(Const(10)), v3)

val p5 = TopLevel(v4)
assert(typeOfProgram(p5) == NumType, "Test failed")
passed(2)

//END TEST
```

Output:
```
at typeOf(<console>:171)
  at typeOfProgram(<console>:217)
  ... 44 elided
```

## Problem 2 (20 points): Objects in Scala

### A (12 points)
Consider the class below for an employee record in an organization that has three fields `name`, `eid` (employee id) and `salary`.

1) As given, scala will complain that class employee cannot extend the abstract class People. Add the missing methods to enable the class Employee to properly extend People

2) Write a factory pattern that takes in a formatted string of the form "John Doe,29585,33110" consisting of a string and two numbers, to extract the name and salaries for each employee. You may want to use `scala.util.matching.Regex` [https://www.scala-lang.org/api/2.9.2/scala/util/matching/Regex.html]
If the input fails to match the pattern, an `IllegalArgumentException` must be thrown

Companion objects are covered in scala book Chapter 4.3. Also lookup here: https://alvinalexander.com/scala/factory-pattern-in-scala-design-patterns

<!-- grade_id: cell-f80a491b407a7ab7, points: 0, grade: false, solution: true -->
```python
import scala.util.matching.Regex

abstract class People {
    /* I need the ability to convert every person to a string */
    def toString: String
    /* I need the ability to check if two people are the same */
    def equals (p: People): Boolean
}

/* 
  Note that for an employee to be equal to an object of type People, that object
  must be an employee with matching eid number 
*/

class Employee( val name: String, val eid: Int, var salary: Int) extends People {
    override def toString = name + ", " + eid.toString + ", " + salary.toString
    def equals(p: People): Boolean = {
        if(Employee(p.toString).eid == eid) return true
        return false
    } 
}

object Employee {
    def apply(formattedInput: String): Employee = {
        val regexPattern = """[a-zA-Z\s+]+[^,\s]+""".r
        val pieces = regexPattern.findAllIn(formattedInput).toArray
        if(pieces.size == 0){
            throw new IllegalArgumentException("Invalid input, try to follow proper formatting")
        }
        else{
            val name = pieces(0).trim
            val eid = pieces(1).trim.toInt
            val salary = pieces(2).trim.toInt
        
            new Employee(name, eid, salary)
        }
    }
}
```

Output:
```
$line39.$read$$iw$$iw$Employee$@6673e38f
```

<!-- grade_id: cell-e1ec3bf75d6f6de5, points: 4, grade: true, solution: false -->
```python
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
println(e1.toString)
assert(e1.name == "Jonathan Doe", "Failed")
assert(e1.eid == 12893, "Failed")
assert(e1.salary == 39813, "Failed")
assert(e1.equals(new Employee("whoknows", 12893, 39813)))
passed(4)
//END TEST
```

Output:
```
Jonathan Doe, 12893, 39813

*** Tests Passed (4 points) ***
null
```

<!-- grade_id: cell-9d823f9842fd95a7, points: 4, grade: true, solution: false -->
```python
//BEGIN TEST
try {
    val e2 = Employee("Badly Formatted String, 29812, $3102")
    assert(false, "Failed, the string is badly formatted but your code is OK with it.")
} catch {
    case e => println(s"Expected behavior seen")
}
passed(4)
//END TEST
```

Output:
```
Expected behavior seen

*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-00d1907d97c7da69, points: 4, grade: true, solution: false -->
```python
//BEGIN TEST
val e1 = Employee("Janet Maria Catherine Doe    , 12894   , 121331")
println(e1.toString)
assert(e1.name == "Janet Maria Catherine Doe", "Failed")
assert(e1.eid == 12894, "Failed")
assert(e1.salary == 121331, "Failed")
assert(e1.equals(new Employee("random", 12894, 121332)))
passed(4)
//END TEST
```

Output:
```
Janet Maria Catherine Doe, 12894, 121331

*** Tests Passed (4 points) ***
null
```

### B (8 points)
We wish to add a new class of people to our system. These are "Contractor" with the information associated should involve their name (string), contract ID (Integer) and contract supervisor (Employee).

Write a new class named `Contractor` whose constructor should take in the name of the person (String), their ID number (Integer) and Supervisor (Employee) and extend from People. 

A contractor object is equal to an object p of type People, only if p is actually an instance of contractor, and if their names are the same and they have the same contract ID number.

<!-- grade_id: cell-97ab958367cfbd0a, points: 0, grade: false, solution: true -->
```python
class Contractor(val name: String, val eid: Int, val supervisor: Employee) extends People {
    override def toString: String = {
        name + ", " + eid.toString() + ", " + supervisor.toString
    }
    override def equals(p: People): Boolean = {
        val formattedInput = p.toString 
        val regexPattern = """[a-zA-Z\s+]+[^,\s]+""".r
        val pieces = regexPattern.findAllIn(formattedInput).toArray
        
        val nameCompare = pieces(0).trim 
        val eidCompare = pieces(1).trim.toInt
        
        if(eid == eidCompare && name == nameCompare){true} else{false}
    } 
}
```

Output:
```
defined class Contractor
```

<!-- grade_id: cell-82cc00e175a12efe, points: 4, grade: true, solution: false -->
```python
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
val c1 = new Contractor("John Contract Worker", 12093, e1)
assert( !c1.equals(e1), "Failed")
assert( c1.equals(c1), "Failed")
passed(4)
//END TEST
```

Output:
```
*** Tests Passed (4 points) ***
null
```

<!-- grade_id: cell-62fa2e0b66f58862, points: 4, grade: true, solution: false -->
```python
//BEGIN TEST
val e1 = Employee("Jonathan Doe, 12893, 39813")
val e2 = Employee("Jane Doe, 12894, 43814")
val c1 = new Contractor("John Contract Worker", 12093, e2)
val c2 = new Contractor("John Contract Worker", 12093, e1)
assert( c1.equals(c2), "Failed")
passed(4)
//END TEST
```

Output:
```
*** Tests Passed (4 points) ***
null
```

## Problem 3 (10 points): Traits in Scala

For this problem, we have defined two traits: `NumberOfLegs` that helps us define how many legs a given animal has and `WarmBlooded` that applies to warm blooded animals.

```python
abstract class Animal 

trait NumberOfLegs {
    val nLegs: Int
    def getNumberOfLegs: Int = nLegs
}

trait WarmBlooded extends Animal {
    val bodyTempMaintained: Double
    def getBodyTemp: Double = bodyTempMaintained 
}
```

Output:
```
defined class Animal
defined trait NumberOfLegs
defined trait WarmBlooded
```

### A (5 points)

Define a class `Human` that inherits from `Animal` and mixes in the traits `NumberOfLegs` with `nLegs = 2` and `WarmBlooded` with `bodyTempMaintained = 98`. The class `Human` must take in a parameter called `name` of type `String` and implement a `getName` method without any parameters.

<!-- grade_id: cell-e7a4a9275bd24e24, points: 0, grade: false, solution: true -->
```python
class Human(val name: String) extends Animal with NumberOfLegs with WarmBlooded{
    val humanName = name
    override val nLegs = 2
    override val bodyTempMaintained = 98.0
    def getName: String = humanName 
}
```

Output:
```
defined class Human
```

<!-- grade_id: cell-f9b7ae8902326372, points: 5, grade: true, solution: false -->
```python
//BEGIN TEST
val t1 = new Human("Jane Smith")
assert(t1.getNumberOfLegs == 2, "Your human does not have two legs")
assert(t1.bodyTempMaintained == 98.0, "Your human does not maintain a body temp of 98")
assert(t1.getBodyTemp == 98.0, "Your human's getBodyTemp Function is not working")
assert(t1.getName == "Jane Smith", "Your human's name is not setting correctly")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
null
```

### B (5 points)
Define a class named `Table` that mixes in the trait `NumberOfLegs` with `nLegs = 4`. 

_Note:_ `Table` cannot mixin the trait `WarmBlooded` (make sure you understand why that is as a quick check of concepts).

<!-- grade_id: cell-8145df690b6f6588, points: 0, grade: false, solution: true -->
```python
class Table extends NumberOfLegs{
    override val nLegs = 4
}
```

Output:
```
defined class Table
```

<!-- grade_id: cell-d34b29b062b93894, points: 5, grade: true, solution: false -->
```python
//BEGIN TEST
val tbl = new Table()
assert(tbl.getNumberOfLegs == 4, "A Table must have four legs")
passed(5)
//END TEST
```

Output:
```
*** Tests Passed (5 points) ***
null
```

```python
def foo(x: String) : Int = {
     x.length * 5
}

foo("hello")
```

Output:
```
25
```

```python
255/0
```

Output:
```
java.lang.ArithmeticException: / by zero
  ... 44 elided
```

```python
def myRecFun_cps[T](x1: Int, x2: String, k: String=> T): T
myRecFun_cps(10, 12, (v) => v)
```

Output:
```
required: String
       myRecFun_cps(10, 12, (v) => v)
                        ^
```

```python
def mkList(x: List[Int], g: Int => List[Int] ): List[Int] =  x.flatMap(g)
val seq = List[4, 5, 6, 7]
val seq1 = List[4, 5, 6, 7,8]
mkList(seq,seq1);
```

Output:
```
<console>:4: error: ']' expected but ';' found.
mkList(seq,seq1);
^
```

```python
def decrypt(x: String): String = {
    
    val freq = Array(0.080, 0.015, 0.030, 0.040, 0.130, 0.020, 0.015, 0.060, 0.065, 0.005, 0.005, 0.035, 0.030, 0.070, 0.080, 0.020, 0.002, 0.065, 0.060, 0.090, 0.030, 0.010, 0.015, 0.005, 0.020, 0.002);
    val let = Array("A","B","C","D","E","F","G","H","I","J","K","L","M","N","O","P","Q","R","S","T","U","V","W","X","Y","Z");
    for(i <- 0 to (freq.length-1)){
        
        print(let(i) + ": " + freq(i) + " (");
        var change = Array(16,4,1,12,23,15,19,11,14,0,5,13,22)
        for(j <- 0 to (change.length-1)){
            if((change(j)+i) > 25)
            {
                change(j) = change(j)+i - 26
            }
            else
            {
                change(j) = change(j)+i
            }
            print(change(j)+",")
        }
        print(") ")
        var tot = ((.05)*freq(change(0))) + ((.05)*freq(change(1))) + ((.05)*freq(change(2))) + ((.05)*freq(change(3))) + ((.05)*freq(change(4))) + ((.15)*freq(change(5))) + ((.05)*freq(change(6))) + ((.05)*freq(change(7))) + ((.05)*freq(change(8))) + ((.05)*freq(change(9))) + ((.05)*freq(change(10))) + ((.2)*freq(change(11))) + ((.15)*freq(change(12))) 
        println(tot)
    }
    return("DONE")
}
decrypt("qeb mxpptloa fp nnwwnnw");
```

Output:
```
A: 0.08 (16,4,1,12,23,15,19,11,14,0,5,13,22,) 0.04360000000000001
B: 0.015 (17,5,2,13,24,16,20,12,15,1,6,14,23,) 0.0328
C: 0.03 (18,6,3,14,25,17,21,13,16,2,7,15,24,) 0.0352
D: 0.04 (19,7,4,15,0,18,22,14,17,3,8,16,25,) 0.04195000000000001
E: 0.13 (20,8,5,16,1,19,23,15,18,4,9,17,0,) 0.0561
F: 0.02 (21,9,6,17,2,20,24,16,19,5,10,18,1,) 0.03185
G: 0.015 (22,10,7,18,3,21,25,17,20,6,11,19,2,) 0.04035
H: 0.06 (23,11,8,19,4,22,0,18,21,7,12,20,3,) 0.042499999999999996
I: 0.065 (24,12,9,20,5,23,1,19,22,8,13,21,4,) 0.040249999999999994
J: 0.005 (25,13,10,21,6,24,2,20,23,9,14,22,5,) 0.021599999999999998
K: 0.005 (0,14,11,22,7,25,3,21,24,10,15,23,6,) 0.021800000000000003
L: 0.035 (1,15,12,23,8,0,4,22,25,11,16,24,7,) 0.04095000000000001
M: 0.03 (2,16,13,24,9,1,5,23,0,12,17,25,8,) 0.028750000000000005
N: 0.07 (3,17,14,25,10,2,6,24,1,13,18,0,9,) 0.039850000000000003
O: 0.08 (4,18,15,0,11,3,7,25,2,14,19,1,10,) 0.03910000000000001
P: 0.02 (5,19,16,1,12,4,8,0,3,15,20,2,11,) 0.05035
Q: 0.002 (6,20,17,2,13,
```

```python

```

```python

```